# Problem 6

In this problem, you will write code to "parse" a restricted form of SQL queries. These exercises are about string processing and regular expressions. There are five (5) exercises, numbered 0-4, which are worth a total of ten (10) points.

```python
from IPython.display import display

import re
import pandas as pd

# Random number generation for generating test cases:
from random import randrange, randint, choice, sample
```

**Background: SQL review.** Suppose you have two SQL tables, `OneTable` and `AnotherTable`, and you wish to perform an inner-join on that links a column named `ColA` in `OneTable` with `ColB` in `AnotherTable`. Recall that one simple way to do that in SQL would be:

```SQL
    SELECT * FROM OneTable, AnotherTable
        WHERE OneTable.ColA = AnotherTable.ColB
```

Or, consider the following more complex example. Suppose you have an additional table named `YetAThird` and you wish to extend the inner-join to include matches between its column, `ColC`, and a second column from `AnotherTable` named `ColB2`:

```SQL
    SELECT * FROM OneTable, AnotherTable, YetAThird
        WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2 = YetAThird.ColC
```

**Exercise 0** (2 points). Suppose you are given a string containing an SQL query in the restricted form,

```SQL
    SELECT * FROM [tbls] WHERE [conds]
```

Implement the function, **`split_simple_join(q)`**, so that it takes a query string **`q`** in the form shown above and **returns a pair of substrings** corresponding to `[tbls]` and `[conds]`.

For example, if

```python
    q == """SELECT * FROM OneTable, AnotherTable, YetAThird
              WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC"""
```

then

```python
    split_simple_join(q) == ("OneTable, AnotherTable, YetAThird",
                             "OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC")
```

**IMPORTANT NOTE!** In this problem, you only need to return the substring between `FROM` and `WHERE` and the one after `WHERE`. You will extract the table names and conditions later on, below.

You should make the following assumptions:

* The input string `q` contains exactly one such query, with no nesting of queries (e.g., no instances of `"SELECT * FROM (SELECT ...)"`). However, the query may (or may not) be a multiline string as shown in the example. (Treat newlines as whitespace.)
* Your function should ignore any leading or trailing whitespace around the SQL keywords, e.g., `SELECT`, `FROM`, and `WHERE`.
* The substring between `SELECT` and `FROM` will be any amount of whitespace, followed by an asterisk (`*`). 
* You should **not** treat the SQL keywords in a case-sensitive way; for example, you would regard `SELECT`, `select`, and `sElEct` as the same. However, do **not** change or ignore the case of the non-SQL keywords.
* The `[tbls]` substring contains only a simple list of table names and no other substrings that might be interpreted as SQL keywords.
* The `[conds]` substring contains only table and column names (e.g., `OneTable.ColA`), the equal sign, the `AND` SQL keyword, and whitespace, but no other SQL keywords or symbols.

> Assuming you are using regular expressions for this problem, recall that you can pass [`re.VERBOSE`](https://docs.python.org/3/library/re.html#re.VERBOSE) when writing a multiline regex pattern.

<!-- grade_id: split_simple_join, points: 0, grade: false, solution: true -->
```python
def split_simple_join(q):
    assert type(q) is str
    ###

    re_s = re.compile(r'''
                        FROM[\s]*?(?P<tbls>[\w+\,\s\S]+)[\s\S]*?WHERE[\s\S]*?
                        (?P<conds>[\w\.\s\=]+)
                        ''',re.VERBOSE)

    result_list = re_s.findall(q, re.IGNORECASE)

    return tuple([i.strip() for i in result_list[0]])
    ###

def split_simple_join(q):
    assert type(q) is str
    ###

    match = re.match(r'''\s*SELECT\s+\*\s+FROM\s+(.*)\s+WHERE\s+(.*)\s*''',q, re.IGNORECASE)

    assert match is not None
    return match.groups()[0], match.groups()[1]
    ###
    
# Demo
q_demo = """SELECT * FROM OneTable, AnotherTable, YetAThird
              WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC"""
print(split_simple_join(q_demo))
```

Output:
```
('OneTable, AnotherTable, YetAThird', 'OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC')
```

<!-- grade_id: split_simple_join_test1, points: 1, grade: true, solution: false -->
```python
# Test cell: `split_simple_join_test1`

assert split_simple_join(q_demo) == \
           ('OneTable, AnotherTable, YetAThird',
            'OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC')
print("\n(Passed!)")
```

Output:
```
(Passed!)
```

<!-- grade_id: split_simple_join_test2, points: 1, grade: true, solution: false -->
```python
# Test cell: `split_simple_join_test2`

__SQL = {'SELECT', 'FROM', 'WHERE'} # SQL keywords

# Different character classes
__ALPHA = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ_'
__ALPHA_VARS = __ALPHA + '_'
__NUM = '0123456789'
__ALPHA_NUM_VARS = __ALPHA_VARS + __NUM
__DOT = '.'
__SPACES = ' \t\n'
__EQ = '='
__ALL = __ALPHA_NUM_VARS + __DOT + __SPACES + __EQ

def flip_coin():
    from random import choice
    return choice([True, False])

def rand_str(k_min, k_max, alphabet):
    """
    Returns a random string of `k` letters chosen uniformly from
    from `alphabet` with replacement, where `k_min <= k <= k_max`
    where `k` is also chosen uniformly at random.
    """
    assert k_max >= k_min >= 0
    k = k_min + randint(0, k_max-k_min)
    return ''.join([choice(alphabet) for _ in range(k)])

def rand_spaces(k_max, k_min=0):
    assert k_max >= k_min >= 0
    return rand_str(k_min, k_max, __SPACES)

def rand_case(s):
    """Randomly changes the case of each letter in a string."""
    return ''.join([c.upper() if flip_coin() else c.lower() for c in s])

def rand_var(k_min=1, k_max=10):
    assert k_max >= k_min >= 1
    s = choice(__ALPHA_VARS)
    s += rand_str(k_min-1, k_max-1, __ALPHA_NUM_VARS)
    if s.upper() in __SQL: # Don't generate keywords -- try again
        return rand_var(k_min, k_max)
    return s

def rand_vars(k_min, k_max):
    V = set()
    dk = randint(0, k_max-k_min)
    for k in range(k_min + dk + 1):
        v = rand_var()
        V.add(v)
    return V

def rand_table(max_cols):
    table = rand_var()
    columns = rand_vars(1, max_cols)
    return (table, columns)

def rand_tables(k_min, k_max, c_max=4):
    assert k_max >= k_min >= 1
    num_tables = k_min + randint(0, k_max-k_min)
    tables = {}
    for _ in range(num_tables):
        table, columns = rand_table(c_max)
        while table in tables:
            table, columns = rand_table(v_max, c_max)
        tables = {**tables, **{table: columns}}
    return tables

def rand_field(table_name, col_names):
    return table_name + "." + choice(list(col_names))

def rand_cond(tables):
    assert type(tables) is dict
    assert len(tables) >= 2
    a, b = sample(tables.keys(), 2)
    a_col = rand_field(a, tables[a])
    b_col = rand_field(b, tables[b])
    return (a_col, b_col)

def rand_cond_set(tables, k_min, k_max):
    k = k_min + randint(0, k_max-k_min)
    conds = set()
    while len(conds) < k:
        (a, b) = rand_cond(tables)
        if ((a, b) not in conds) and ((b, a) not in conds):
            conds.add((a, b))
    return conds

def cond_set_to_substrs(conds, max_wspad=4):
    substrs = []
    for a, b in conds:
        s = "{}{}{}{}{}".format(a,
                                rand_spaces(max_wspad),
                                __EQ,
                                rand_spaces(max_wspad),
                                b)
        substrs.append(s)
    return substrs

def substrs_to_str(substrings, sep=',', max_wspad=4):
    s_final = ''
    for k, s in enumerate(substrings):
        if k > 0:
            s_final += rand_spaces(max_wspad) + sep + rand_spaces(max_wspad)
        s_final += s
    return s_final

def rand_query_ans(max_tables, max_conds, max_wspad=4):
    tables = rand_tables(2, max_tables)
    cond_set = rand_cond_set(tables, 1, 4)
    return tables, cond_set

def pad_1(max_wspad):
    return rand_spaces(max(1, max_wspad), 1)

def form_select(max_wspad=4):
    return pad_1(max_wspad) + rand_case("SELECT") + pad_1(max_wspad) + "*"

def form_from(tables, max_wspad=4):
    from_ans = substrs_to_str(list(tables.keys()), sep=',', max_wspad=max_wspad)
    return pad_1(max_wspad) + rand_case("FROM") + pad_1(max_wspad) + from_ans, from_ans

def form_where(cond_set, max_wspad=4):
    cond_substrs = cond_set_to_substrs(cond_set)
    cond_ans = substrs_to_str(cond_substrs, sep=' AND ', max_wspad=max_wspad)
    return pad_1(max_wspad) + rand_case("WHERE") + pad_1(max_wspad) + cond_ans, cond_ans

def form_query_str(tables, cond_set, max_wspad=4):
    select_clause = form_select(max_wspad)
    from_clause, from_ans = form_from(tables, max_wspad)
    where_clause, cond_ans = form_where(cond_set, max_wspad)
    query = select_clause + from_clause + where_clause
    return query, from_ans, cond_ans

def split_simple_join_battery(num_tests, max_wspad=4):
    for k in range(num_tests):
        tables, cond_set = rand_query_ans(5, 5, max_wspad)
        qstmt, from_ans, where_ans = form_query_str(tables, cond_set, max_wspad)
        print("=== Test Statement {} ===\n'''{}'''\n".format(k, qstmt))
        print("True 'FROM' clause substring: '''{}'''\n".format(from_ans))
        print("True 'WHERE' clause substring: '''{}'''\n".format(where_ans))
    
split_simple_join_battery(5, 3)

print("\n(Passed!)")
```

Output:
```
=== Test Statement 0 ===
''' 	SELEct	* 
FrOM	Muc_kFyJ  	,
Nd,
 
vK,
	Qx	whEre
Muc_kFyJ.qYLHirnB
	=
Nd.qqcSj AND vK.ec4UtoRd8=		Muc_kFyJ.qYLHirnB
	  AND 	
Nd.qqcSj	= 	 Muc_kFyJ.wrv	   AND 
	
vK.amxqGroj	

=	Nd.qqcSj'''

True 'FROM' clause substring: '''Muc_kFyJ  	,
Nd,
 
vK,
	Qx'''

True 'WHERE' clause substring: '''Muc_kFyJ.qYLHirnB
	=
Nd.qqcSj AND vK.ec4UtoRd8=		Muc_kFyJ.qYLHirnB
	  AND 	
Nd.qqcSj	= 	 Muc_kFyJ.wrv	   AND 
	
vK.amxqGroj	

=	Nd.qqcSj'''

=== Test Statement 1 ===
''' SEleCt 	* fROm 		dl,xE3TE6	
,   rb_U5k 	,

jWlOEDh,_MpB
WHeRE		 jWlOEDh.QZjFHi  
	= 
xE3TE6.JqMayTqKad'''

True 'FROM' clause substring: '''dl,xE3TE6	
,   rb_U5k 	,

jWlOEDh,_MpB'''

True 'WHERE' clause substring: '''jWlOEDh.QZjFHi  
	= 
xE3TE6.JqMayTqKad'''

=== Test Statement 2 ===
'''	SeLeCT	
 *  froM
_n, j ,   isn1hC 
 wheRE  isn1hC.QWtEKjMmA

=
 
	j.Av'''

True 'FROM' clause substring: '''_n, j ,   isn1hC'''

True 'WHERE' clause substring: '''isn1hC.QWtEKjMmA

=
 
	j.Av'''

=== Test Statement 3 ===
'''	
```

**Variable names.** For this problem, let a valid _variable name_ be a sequence of alphanumeric characters or underscores, where the very first character _cannot_ be a number. For example,

    some_variable
    __another_VariAble
    Yet_a_3rd_var
    _A_CSE_6040_inspired_var
    
are all valid variable names, whereas the following are not.

    123var_is_bad
    0_is_not_good_either
    4goodnessSakeStopItAlready

**Exercise 1** (2 points). Implement a function, **`is_var(s)`**, that checks whether a valid variable name. That is, it should return `True` if and only if `s` is valid according to the above definition. Your function should ignore any leading or trailing spaces in `s`.

For example:

```python
    assert is_var("__another_VariAble")
    assert not is_var("0_is_not_good_either")
    assert is_var("   foo")
    assert is_var("_A_CSE_6040_inspired_var   ")
    assert not is_var("#getMe2")
    assert is_var("   Yet_a_3rd_var  ")
    assert not is_var("123var_is_bad")
    assert not is_var("  A.okay")
```

<!-- grade_id: is_var, points: 0, grade: false, solution: true -->
```python
def is_var(s):
    assert type(s) is str
    ###
#     s = '0_is_not_good_either'
    re_s = re.compile(r'''
                            [\W]

                            ''',re.VERBOSE)
    if s.strip()[0].isdigit():
        return   False
    else:

        result_list = re_s.findall(s.strip())
        return len(result_list) == 0

def is_var(s):
    assert type(s) is str
    ###
    matches = re.match(r'''^\s*  
                        ([a-zA-Z_]
                        \w*)
                        \s*$                     
                            ''', s,re.VERBOSE)
    return matches is not None
    ###

```

```python
is_var('-----*lF0OOAGW------*')


```

Output:
```
False
```

<!-- grade_id: is_var_test0, points: 1, grade: true, solution: false -->
```python
# Test cell, part 1: `is_var_test0`

assert is_var("__another_VariAble")
assert not is_var("0_is_not_good_either")
assert is_var("   foo")
assert is_var("_A_CSE_6040_inspired_var   ")
assert not is_var("#getMe2")
assert is_var("   Yet_a_3rd_var  ")
assert not is_var("123var_is_bad")
assert not is_var("  A.okay")

print("\n(Passed part 1 of 2.)")
```

Output:
```
(Passed part 1 of 2.)
```

<!-- grade_id: is_var_test1, points: 1, grade: true, solution: false -->
```python
# Test cell: `is_var_test2`

for v in rand_vars(20, 30):
    ans = flip_coin()
    if not ans:
        v = choice(__NUM) + v
    v = rand_spaces(3) + v + rand_spaces(3)
    your_ans = is_var(v)
    print(v)
    assert your_ans == ans, "is_var('{}') == {} instead of {}.".format(v, your_ans, ans)
    
print("\n(Passed part 2 of 2.)")
```

Output:
```
7TJ
		3mRwHL
	
 h		
	 	k   
		6kk

	WB0Kl

Dnf

	 3g5i7
		XcX7x1ff		

	9NVtldC



 0f8KSfO2F
 v9SyY
 
d 

HGB880Fm 
	
 
	Dvq 
6syfrovy 	 

 6aTc_x9Jx	
 
Kz_Vl5iJu

 
Ix 
	

 sr	

	
8r9
	 
3y6v2AzHNTw


	7l8 
  
Jxk0TUz
	
Br

7eQ  

	

5rHIMA1qhC5	
9UFzAgN	

	holqW	 

(Passed part 2 of 2.)
```

**Column variables.** A _column variable_ consists of two valid variable names separated by a single period. For example,

    A.okay
    a32X844._387b
    __C__.B3am
    
are all examples of column variables: in each case, the substrings to the left and right of the period are valid variables.

**Exercise 2** (1 point). Implement a function, **`is_col(s)`**, so that it returns `True` if and only if **`s`** is a column variable, per the definition above.

For example:

```python
    assert is_col("A.okay")
    assert is_col("a32X844._387b")
    assert is_col("__C__.B3am")
    assert not is_col("123.abc")
    assert not is_col("abc.123")
```

As with Exercise 1, your function should ignore any leading or trailing spaces.

<!-- grade_id: is_col, points: 0, grade: false, solution: true -->
```python
def is_col(s):
    ###
    
# s = "\t \n4750.55\t\n"
    if '.' in s:
        m = re.split(r"\.", s)    

        return  is_var(m[0]) & is_var(m[1])
    else:
        return False
def is_col(s):
    matches = re.match(r'''\s*
                    ([a-zA-Z_]\w*)
                    \.
                    ([a-zA-Z_]\w*)
                    \s*    ''', s,  re.VERBOSE)
    return matches is not None
    ###

```

<!-- grade_id: is_col_test0, points: 1, grade: true, solution: false -->
```python
# Test cell: `is_col_test0`

assert is_col("A.okay")
assert is_col("a32X844._387b")
assert is_col("__C__.B3am")
assert not is_col("123.abc")
assert not is_col("abc.123")

print("\n(Passed part 1.)")
```

Output:
```
(Passed part 1.)
```

<!-- grade_id: is_col_test1, points: 1, grade: true, solution: false -->
```python
# Test cell: `is_col_test1`

def test_is_col_1():
    a = rand_var()
    assert not is_col(a), "is_col('{}') == {} instead of {}.".format(a, is_col(a), False)
    a_valid = flip_coin()
    if not a_valid:
        a = rand_str(1, 5, __NUM)
    return a, a_valid

for _ in range(20):
    a, a_valid = test_is_col_1()
    b, b_valid = test_is_col_1()
    ans = a_valid and b_valid
    
    c = "{}{}.{}{}".format(rand_spaces(3), a, b, rand_spaces(3))
    your_ans = is_col(c)
    print("==> is_col('{}') == {}".format(c, your_ans))
    assert your_ans == ans

print("\n(Passed part 2.)")
```

Output:
```
==> is_col('	 	3.073		 ') == False
==> is_col('
90892.7 
 ') == False
==> is_col('

kCIqurP.8
		') == False
==> is_col(' fLd7zx.iAnlLI3r8 ') == True
==> is_col(' 	
_ky44oHn5.aimCRtU
 
') == True
==> is_col('3.862  	') == False
==> is_col('		78.52836') == False
==> is_col(' 	2.Yyzx5Z') == False
==> is_col('		fM.4') == False
==> is_col('77.74274 ') == False
==> is_col(' 
le2b9gWdD.693

	') == False
==> is_col('
	4.096') == False
==> is_col('		MX_pf_qV_.237

') == False
==> is_col('81.LP6cYify_g') == False
==> is_col('

Ri0G8yqW._V
 ') == True
==> is_col(' 
 GGQMR_BTr6.RE8o	
') == True
==> is_col('87.91706	
') == False
==> is_col('	
37.9		') == False
==> is_col('

bIlJxb7.bhybCqFx88
	') == True
==> is_col('
7659.02') == False

(Passed part 2.)
```

**Equality strings.** An _equality string_ is a string of the form,

    A.x = B.y

where `A.x` and `B.y` are _column variable_ names and `=` is an equals sign. There may be any amount of whitespace---including none---before or after each variable and the equals sign.

**Exercise 3** (2 points). Implement the function, **`extract_eqcols(s)`**, below. Given an input string **`s`**, if it is an equality string, your function should return a pair `(u, v)`, where `u` and `v` are the two column variables in the equality string. For example:

```python
    assert extract_eqcols("F3b._xyz =AB0_.def") == ("F3b._xyz", "AB0_.def")
```

If `s` is not a valid equality string, then your function should return `None`.

<!-- grade_id: extract_eqcols, points: 0, grade: false, solution: true -->
```python
def extract_eqcols(s):
    ###

#     s = "F3b._xyz =AB0_.def"
    
    m = re.match(r'(\w+\.\w+)\s*?\=\s*?(\w+\.\w+)',s)
    if '=' in s:
        if is_col(m.groups()[0]) & is_col(m.groups()[1]):
            return m.groups()
        ###
    
print(extract_eqcols("F3b._xyz =AB0_.def"))
```

Output:
```
('F3b._xyz', 'AB0_.def')
```

<!-- grade_id: extract_eqcols0, points: 1, grade: true, solution: false -->
```python
# Test cell: `extract_eqcols0`

assert extract_eqcols("F3b._xyz =AB0_.def") == ("F3b._xyz", "AB0_.def")
assert extract_eqcols("0F3b._xyz =AB0_.def") is None

print("\n(Passed part 1 of 2.)")
```

Output:
```
(Passed part 1 of 2.)
```

<!-- grade_id: extract_eqcols1, points: 1, grade: true, solution: false -->
```python
# Test cell: `extract_eqcols1`

for _ in range(5):
    _, cond_set = rand_query_ans(2, 10, 5)
    for a, b in cond_set:
        s = a + rand_spaces(3) + __EQ + rand_spaces(3) + b
        print("==> Processing:\n'''{}'''\n".format(s))
        ans = extract_eqcols(s)
        print("    *** Found: {} ***".format(ans))
        assert ans is not None, "Did not detect an equality string where there was one!"
        assert ans[0] == a and ans[1] == b, "Returned {} instead of ({}, {})".format(ans, a, b)

print("\n(Passed part 2 of 2.)")
```

Output:
```
==> Processing:
'''LWI.z=
		YxXO9wTs6N.EVZ3TtLC'''

    *** Found: ('LWI.z', 'YxXO9wTs6N.EVZ3TtLC') ***
==> Processing:
'''YxXO9wTs6N.XpI_dFHEF3
	
=
LWI.z'''

    *** Found: ('YxXO9wTs6N.XpI_dFHEF3', 'LWI.z') ***
==> Processing:
'''LWI.C3bIuFxb7=YxXO9wTs6N.XpI_dFHEF3'''

    *** Found: ('LWI.C3bIuFxb7', 'YxXO9wTs6N.XpI_dFHEF3') ***
==> Processing:
'''LWI.z 		=YxXO9wTs6N.Yn3Bbd'''

    *** Found: ('LWI.z', 'YxXO9wTs6N.Yn3Bbd') ***
==> Processing:
'''efLp8.k=
uQex6l.d'''

    *** Found: ('efLp8.k', 'uQex6l.d') ***
==> Processing:
'''efLp8.k
=

uQex6l.Q2u_'''

    *** Found: ('efLp8.k', 'uQex6l.Q2u_') ***
==> Processing:
'''efLp8.k
=	uQex6l.fI5q'''

    *** Found: ('efLp8.k', 'uQex6l.fI5q') ***
==> Processing:
'''q.ZPtD5s  =

mZg.AhTiQ'''

    *** Found: ('q.ZPtD5s', 'mZg.AhTiQ') ***
==> Processing:
'''mZg.Wy5g=q.BXR'''

    *** Found: ('mZg.Wy5g', 'q.BXR') ***
==> Processing:
'''mZg.Wy5g=q.ZPtD5s'''

    *** Found: ('mZg.Wy5g', 'q.ZPtD5s') ***
==> Processing:
'''mZg.Wy5g =q.xldX34anLo'''
```

**Exercise 4** (2 points). Given an SQL query in the restricted form described above, write a function that extracts all of the join conditions from the `WHERE` clause. Name this fuction, **`extract_join_conds(q)`**, where `q` is the query string. It should return a list of pairs, where each pair `(a, b)` is the name of the left- and right-hand sides in one of these conditions.

For example, suppose:

```python
    q == """SELECT * FROM OneTable, AnotherTable, YetAThird
              WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC"""
```

Notice that the `WHERE` clause contains two conditions: `OneTable.ColA = AnotherTable.ColB` and `AnotherTable.ColB2=YetAThird.ColC`. Therefore, your function should return a list of two pairs,
as follows:

```python
    extract_join_conds(q) == [("OneTable.ColA", "AnotherTable.ColB"),
                              ("AnotherTable.ColB2", "YetAThird.ColC")]
```

<!-- grade_id: extract_join_conds, points: 0, grade: false, solution: true -->
```python
def extract_join_conds(q):
    ###
# q = """SELECT * FROM OneTable, AnotherTable, YetAThird
#               WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC"""
    result_list = re.findall(r'''WHERE[\s]*?(?P<tbls>[\w+\,\s\S]+)[\s\S]*?''', q,  re.IGNORECASE)
    
    slist = result_list[0].split('AND')
    li = []
    for s in slist:
        st = extract_eqcols(s.strip())
        li.append(st)
    
    return li
    ###
    
print("==> Query:\n\t'{}'\n".format(q_demo))
print("==> Results:\n{}".format(extract_join_conds(q_demo)))
```

Output:
```
==> Query:
	'SELECT * FROM OneTable, AnotherTable, YetAThird
              WHERE OneTable.ColA = AnotherTable.ColB AND AnotherTable.ColB2=YetAThird.ColC'

==> Results:
[('OneTable.ColA', 'AnotherTable.ColB'), ('AnotherTable.ColB2', 'YetAThird.ColC')]
```

<!-- grade_id: extract_join_conds_test, points: 2, grade: true, solution: false -->
```python
# Test cell: `extract_join_conds_test`

def test_extract_join_conds_1():
    tables, cond_set = rand_query_ans(5, 5, 0)
    qstmt, _, _ = form_query_str(tables, cond_set, 0)
    qstmt = re.sub("[\n\t]", " ", qstmt)
    print("=== {} ===\n".format(qstmt))
    print("  True solution: {}\n".format(cond_set))
    your_conds = extract_join_conds(qstmt)
    print("  Your solution: {}\n".format(your_conds))
    assert set(your_conds) == cond_set, "*** Mismatch? ***"
    
for _ in range(10):
    test_extract_join_conds_1()
    
print("\n(Passed!)")
```

Output:
```
===  SeleCt * FROM p,tk,So_seixm WHerE tk.DKq= p.eV5TV AND So_seixm.U4sec=p.eV5TV AND So_seixm.SIcf=p.YHaQYW2o ===

  True solution: {('tk.DKq', 'p.eV5TV'), ('So_seixm.U4sec', 'p.eV5TV'), ('So_seixm.SIcf', 'p.YHaQYW2o')}

  Your solution: [('tk.DKq', 'p.eV5TV'), ('So_seixm.U4sec', 'p.eV5TV'), ('So_seixm.SIcf', 'p.YHaQYW2o')]

===  SELEct * From te,dA6S whErE dA6S.xR=    te.ue3bu ===

  True solution: {('dA6S.xR', 'te.ue3bu')}

  Your solution: [('dA6S.xR', 'te.ue3bu')]

===  SElecT * FrOm q6RhCIgd,S,NfMT3pQ,A wHeRe S.F=A.Tq9Q ===

  True solution: {('S.F', 'A.Tq9Q')}

  Your solution: [('S.F', 'A.Tq9Q')]

===  seLeCt * FRoM KF,DQ_zYwH,GY_,Cl_hZtttg,PSfbL WHeRe Cl_hZtttg.oWS7Fb   =    DQ_zYwH.JCWi8_ ===

  True solution: {('Cl_hZtttg.oWS7Fb', 'DQ_zYwH.JCWi8_')}

  Your solution: [('Cl_hZtttg.oWS7Fb', 'DQ_zYwH.JCWi8_')]

===  sELECt * fRoM DlwKgt,QyrP,J5k,wNG wHERe J5k.xO9C5LFoV   =  QyrP.F9Nju AND wNG.uGf8sa   =J5k.dWzJFOp6fq AND wNG.oZTiZ4dlJ  =   J5k.xO9C5LFoV ===

  True solution: {(
```

**Fin!** This marks the end of this problem. Don't forget to submit it to get credit.